# CS295/395: Secure Distributed Computation
## Homework 1

```python
# Useful imports and utility functions

import galois
GF = galois.GF(2**31-1)

def make_additive_shares(n, x):
    """Create additive secret shares of the input, so that the shares look random but add up to the input.
    n: number of shares to create
    x: input value"""
    first_shares = GF.Random(n-1)
    last_share = x - first_shares.sum()
    return GF(list(first_shares) + [last_share])
```

## Question 1 (20 points)

Imagine you (the "client") want to periodically submit encrypted numbers to two "servers" for storage. After some time, you'd like to have the servers add the numbers up and return the encrypted sum, so that you don't need to compute the sum yourself.

Implement a system for *outsourced computation* of the sum of a list of numbers to solve this problem by filling in the functions below. Your solution should simulate who holds what data by updating the `server_state` and `client_state` lists.

You can assume that at least one server is **honest**, and the other is **semi-honest**. You don't know which is which ahead of time. Your solution should prevent the semi-honest server from learning any of the client's inputs.

<!-- grade_id: cell-ee91ecf45af4acb4, points: 0, grade: false, solution: true -->
```python
# Initialize server and client state
server1_state = []
server2_state = []

# Functions to simulate sending values to each of the servers
def send_to_server1(x):
    server1_state.append(x)

def send_to_server2(x):
    server2_state.append(x)

def init():
    '''Initialize the server and client state'''
    global server1_state
    global server2_state
    global client_state

    server1_state = []
    server2_state = []

def store_number(x):
    '''Store a new input number x.'''
    # YOUR CODE HERE
    raise NotImplementedError()

def run_server1_computation():
    '''Run the server-side computation for server 1, returning its result.'''
    # YOUR CODE HERE
    raise NotImplementedError()
    
def run_server2_computation():
    '''Run the server-side computation for server 2, returning its result.'''
    # YOUR CODE HERE
    raise NotImplementedError()

def decrypt_result(server1_result, server2_result):
    '''Run the client-side decryption process, returning the plaintext result of the
    server-side computation. Your solution should use the `client_state` to perform
    the decryption.'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3135ed9e0b939c73, points: 20, grade: true, solution: false -->
```python
# TEST CASE for problem 1

# A test list of numbers
number_stream = [3,4,5,1,2,3,5,1,2,3,5,5,5,2,3,1,1,5,3,4,5]

# Initialize client & server state
init()

# Handle each number in the stream
for n in number_stream:
    store_number(n)

# Run the server computation and return it to the client
server1_result = run_server1_computation()
server2_result = run_server2_computation()

# Decrypt the result of the server computation
decrypted_result = decrypt_result(server1_result, server2_result)

# It should be the same as the sum of the original number stream
print(decrypted_result)
print(sum(number_stream))
assert decrypted_result == sum(number_stream)
```

## Question 2 (5 points)

Why do we choose $p = 2^{31}-1$? Does it matter that $p$ is prime? What could happen if we choose $p$ too small, and the stream of numbers ends up being very large?

<!-- grade_id: cell-cc8628ba7163f0e6, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (5 points)

This system handles only unsigned (non-negative) integers. In 3-5 sentences, describe how we could extend it to handle *signed* integers?

<!-- grade_id: cell-21bb46fb153b3eca, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE