# CS295/395: Secure Distributed Computation
## Homework 10

```python
# Imports and definitions
import numpy as np
from collections import defaultdict
from collections import namedtuple
import urllib.request
import hashlib
from nacl.signing import SigningKey
```

The structure below implements simplified transactions in the style of Bitcoin. Reference [this page](https://en.bitcoin.it/wiki/Transaction) for more information about the structure of transactions.

```python
HashPointer = namedtuple('HashPointer', ['hash', 'pointer'])
Transaction = namedtuple('Transaction', ['inputs', 'outputs'])

Input = namedtuple('Input', ['previous_tx', 'index', 'public_key'])
SignedInput = namedtuple('SignedInput', ['input', 'signed_input'])

Output = namedtuple('Output', ['public_key_hash', 'value'])

DIFFICULTY = int(2**(32 * 8)/10000)

class Block:
    def __init__(self, transactions, prev, nonce, pubkey_hash):
        self.transactions = transactions
        self.prev = prev
        self.nonce = nonce
        self.pubkey_hash = pubkey_hash
    
    def __repr__(self):
        return f'\nBlock(\n transaction: {self.transactions},\n nonce: {self.nonce},\n prev: {self.prev})'

def bytes_value(v):
    return bytes(str(v), encoding='utf-8')

def hash_value(v):
    return hashlib.sha256(bytes_value(v)).hexdigest()
```

## Question 1 (20 points)

Implement `mine_for_block`, a function to add a block to the blockchain. It should ensure that the hash of the blockchain is less than or equal to `DIFFICULTY`.

<!-- grade_id: cell-0f9e6fbf3cdf909c, points: 0, grade: false, solution: true -->
```python
def add_block(transactions, blockchain, nonce, pubkey_hash):
    prev_hash = hash_value(blockchain)
    prev = HashPointer(prev_hash, blockchain)
    new_block = Block(transactions, prev, nonce, pubkey_hash)
    
    return new_block, hash_value(new_block)

def mine_for_block(transactions, pubkey_hash, blockchain):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ab7d66775a7f646b, points: 10, grade: true, solution: false -->
```python
# TEST CASE
key_bob = SigningKey.generate()
key_alice = SigningKey.generate()
pubkey_bob = key_bob.verify_key
pubkey_alice = key_alice.verify_key

# Bob's going to mine some coins
i1 = Input('COINBASE', 0, pubkey_bob)
i1_signed = SignedInput(i1, key_bob.sign(bytes_value(i1)))
o1 = Output(hash_value(pubkey_bob), 10.0)
tx1 = Transaction([i1_signed], [o1])

# And pay some to Alice
i2 = Input(hash_value(tx1), 0, pubkey_bob)
i2_signed = SignedInput(i2, key_bob.sign(bytes_value(i2)))
o2 = Output(hash_value(pubkey_bob), 9.0)
o3 = Output(hash_value(pubkey_alice), 1.0)
tx2 = Transaction([i2_signed], [o2, o3])

# Mine the block
b1, b1_hash = mine_for_block([tx2, tx1], hash_value(pubkey_bob), None)

# Alice pays some money back to Bob
i3 = Input(hash_value(tx2), 1, pubkey_alice)
i3_signed = SignedInput(i3, key_alice.sign(bytes_value(i3)))
o4 = Output(hash_value(pubkey_alice), 0.5)
o5 = Output(hash_value(pubkey_bob), 0.5)
tx3 = Transaction([i3_signed], [o4, o5])
b2, b2_hash = mine_for_block([tx3], hash_value(pubkey_bob), b1)
```

## Question 2 (20 points)

Implement `validate_blockchain`, a function to validate the blockchain. It should check that:

1. The structure of the blockchain is valid
2. Each block satisfies the difficulty requirement
3. Each transaction is valid

The function should throw an error if the blockchain is not valid, and return `True` if it is valid.

<!-- grade_id: cell-d2158ae8d840ee8c, points: 0, grade: false, solution: true -->
```python
def check_blockchain(blockchain, expected_hash):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-71747a299a955eaf, points: 20, grade: true, solution: false -->
```python
# TEST CASE
assert check_blockchain(b2, b2_hash)
```

## Question 3 (20 points)

In 3-5 sentences, describe what would be required to modify this scheme to support *proof-of-stake mining*, in the style of Peercoin. Reference Section 8.5 in [the textbook](https://d28rh4a8wq0iu5.cloudfront.net/bitcointech/readings/princeton_bitcoin_book.pdf).

<!-- grade_id: cell-c5d7dc947be1718f, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE