# CS295/395: Secure Distributed Computation
## Homework 1

```python
# Useful imports and utility functions

import numpy as np

def plusFE(p, a, b):
    return (a + b) % p
    
def multFE(p, a, b):
    return (a * b) % p

def sumFE(p, xs):
    total = 0
    for x in xs:
        total = plusFE(p, x, total)
    return total
```

## Question 1 (20 points)

Imagine you (the "client") want to periodically submit encrypted numbers to an untrusted "server" for storage. After some time, you'd like to have the server add the numbers up and return the encrypted sum, so that you don't need to compute the sum yourself.

Implement a system for *outsourced computation* of the sum of a list of numbers to solve this problem by filling in the functions below. Your solution should simulate who holds what data by updating the `server_state` and `client_state` lists.

<!-- grade_id: cell-ee91ecf45af4acb4, points: 0, grade: false, solution: true -->
```python
def enc(p, x):
    '''Additively homomorphic encryption for a field element x in GF(p)'''
    # YOUR CODE HERE
    raise NotImplementedError()

def dec(p, k, c):
    '''Additively homomorphic decryption for ciphertext c with key k in GF(p)'''
    # YOUR CODE HERE
    raise NotImplementedError()

# Use GF(2^32) to simulate 32-bit numbers
p = 2**32

# Initialize server and client state
server_state = []
client_state = []

def init():
    '''Initialize the server and client state'''
    global server_state
    global client_state

    server_state = []
    client_state = []

def handle_new_number(x):
    '''Encrypt the field element x and submit it to the server, saving the key on the 
    client. Your function should simulate submission to the server by putting the
    ciphertext in the `server_state` list.'''
    # YOUR CODE HERE
    raise NotImplementedError()

def run_server_computation():
    '''Run the server-side computation, returning its result. The server-side computation 
    should add up all of the numbers in the server's state.'''
    # YOUR CODE HERE
    raise NotImplementedError()

def decrypt_result(server_result):
    '''Run the client-side decryption process, returning the plaintext result of the
    server-side computation. Your solution should use the `client_state` to perform
    the decryption.'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3135ed9e0b939c73, points: 20, grade: true, solution: false -->
```python
# TEST CASE for problem 1

# A test list of numbers
number_stream = [3,4,5,1,2,3,5,1,2,3,5,5,5,2,3,1,1,5,3,4,5]

# Initialize client & server state
init()

# Handle each number in the stream
for n in number_stream:
    handle_new_number(n)

# Run the server computation and return it to the client
server_result = run_server_computation()

# Decrypt the result of the server computation
decrypted_result = decrypt_result(server_result)

# It should be the same as the sum of the original number stream
assert decrypted_result == np.sum(number_stream)
```

## Question 2 (5 points)

Why do we choose $p = 2^{32}$? What could happen if we choose $p$ too small, and the stream of numbers ends up being very large?

<!-- grade_id: cell-cc8628ba7163f0e6, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (5 points)

This system handles only unsigned (non-negative) integers. In 3-5 sentences, describe how we could extend it to handle *signed* integers?

<!-- grade_id: cell-21bb46fb153b3eca, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE