# CS295/395: Secure Distributed Computation
## Homework 4

## Definitions

```python
# Imports and definitions
import numpy as np
from collections import defaultdict
import urllib.request

_PRIME = 2 ** 13 - 1

shamir_lib_url = "https://raw.githubusercontent.com/jnear/cs295-secure-computation/master/utils/shamir.py"

### DANGER: this line is dangerous. Make sure the URL above is correct, and has correct code.
exec(urllib.request.urlopen(shamir_lib_url).read())

class Party:
    """A participant in a multiparty computation protocol."""
    def __init__(self, field_size=_PRIME):
        """Initialize the field size and dictionary to hold received messages."""
        self.field_size = field_size
        self.input = None
        self.output = None
        self.received = defaultdict(list)
    
    def send(self, other, round, msg):
        """Simulate sending a message `msg` to another party `other` during round `round`"""
        other.received[round].append(msg)

    def get_view(self):
        """Returns the view of this party: its input, output, and received messages."""
        return (self.input, self.output, dict(self.received))
```

```python
share_shamir
```

## Question 1 (20 points)

Describe a protocol to multiply three input numbers. The input numbers will be secret-shared according to a $(t,n)$ Shamir secret sharing scheme before the protocol starts, and each party will receive one share of each number. Each party should output *one share of the product*, using a $(t, n)$ Shamir secret sharing scheme (i.e. the threshold for the output should be the same as the threshold for the input).

\begin{equation*}
\textbf{Functionality: Multiply Three Numbers}\\
\fbox{$\mathcal{F}(a, b, c) = a \cdot b \cdot c$}
\end{equation*}

**HINT**: reference the in-class exercise from 9/28/30.

<!-- grade_id: cell-6df68c07c24abf58, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2

Implement your protocol from question 1.

<!-- grade_id: cell-47362684de6912c6, points: 0, grade: false, solution: true -->
```python
class MultThreeParty(Party):
    def round1(self, parties, a_shr, b_shr, c_shr, t):
        self.input = (a_shr, b_shr, c_shr)
        self.c_shr = c_shr # save this one for later
        self.parties = parties
        n = len(parties)
        assert t <= n/2
        
        # YOUR CODE HERE
        raise NotImplementedError()

    def round2(self):
        n = len(self.parties)
        
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def round3(self):
        n = len(self.parties)
        
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-0151a9043c28b815, points: 20, grade: true, solution: false -->
```python
# TEST CASE for question 2

NUM_PARTIES = 6
# (t, n)-Shamir scheme
n = NUM_PARTIES
t = 3

shares1 = share_shamir(t, n, 5)
shares2 = share_shamir(t, n, 6)
shares3 = share_shamir(t, n, 7)


parties = [MultThreeParty() for _ in range(NUM_PARTIES)]

for p,s1,s2,s3 in zip(parties, shares1, shares2, shares3):
    p.round1(parties, s1, s2, s3, t)
for p in parties:
    p.round2()
for p in parties:
    p.round3()
for p in parties:
    print(p.get_view())

output_shares = [p.output for p in parties]
print('Reconstruction, with all shares:', reconstruct_shamir(output_shares))
print('Reconstruction, with 3 shares:', reconstruct_shamir(output_shares[:3]))
print('Reconstruction, with 2 shares:', reconstruct_shamir(output_shares[:2]))

assert reconstruct_shamir(output_shares) == 210
assert reconstruct_shamir(output_shares[:3]) == 210
assert reconstruct_shamir(output_shares[:2]) != 210
```

## Question 3

Describe a protocol to compute the product of a list of $k$ numbers. The input numbers will be secret-shared according to a $(t,n)$ Shamir secret sharing scheme before the protocol starts, and each party will receive a list containing one share of each number. Each party should output *one share of the product*, using a $(t, n)$ Shamir secret sharing scheme (i.e. the threshold for the output should be the same as the threshold for the input).

\begin{equation*}
\textbf{Functionality: Product of $k$ Numbers}\\
\fbox{$\mathcal{F}(x_1, \dots, x_k) = \prod_{i=1}^k x_i$}
\end{equation*}

**HINT**: This problem is intentionally open-ended. The number of rounds of communication between the parties will depend on $k$.

<!-- grade_id: cell-3c058e250a99af4a, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4

Implement your protocol from question 3.

<!-- grade_id: cell-1fe3ab52cdff1e68, points: 0, grade: false, solution: true -->
```python
class MultListParty(Party):
    def round1(self, parties, xs, t):
        self.input = xs
        self.xs = xs
        self.parties = parties
        n = len(parties)
        assert t <= n/2
        
        # YOUR CODE HERE
        raise NotImplementedError()

    def roundn(self, round):
        n = len(self.parties)
        
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def roundk(self, round):
        n = len(self.parties)
        
        # YOUR CODE HERE
        raise NotImplementedError()
```

```python
# Driver function for question 4
# NOTE: you can modify this function, if it helps for your implementation
# You may also want to uncomment pieces of the function for debugging
def run_list_prod():
    NUM_PARTIES = 6
    # (t, n)-Shamir scheme
    n = NUM_PARTIES
    t = 3
    k = 5

    parties = [MultListParty() for _ in range(NUM_PARTIES)]
    nums = [2 for _ in range(k)]
    shares = [share_shamir(t, n, x) for x in nums]
    share_groups = list(zip(*shares))

    for p,xs in zip(parties, share_groups):
        p.round1(parties, xs, t)
        
    for round_num in range(2, k+1):
        for p in parties:
            p.roundn(round_num)
            
    for p in parties:
        p.roundk(k+1)
        
#     for p in parties:
#         print(p.get_view())

    output_shares = [p.output for p in parties]
    
#     print('Reconstruction, with all shares:', reconstruct_shamir(output_shares))
#     print('Reconstruction, with 3 shares:', reconstruct_shamir(output_shares[:3]))
#     print('Reconstruction, with 2 shares:', reconstruct_shamir(output_shares[:2]))
    
    return nums, output_shares
```

<!-- grade_id: cell-c4b396dc2ba93187, points: 20, grade: true, solution: false -->
```python
# TEST CASE for question 4
nums, output_shares = run_list_prod()
assert reconstruct_shamir(output_shares) == np.prod(nums)
assert reconstruct_shamir(output_shares[:3]) == np.prod(nums)
assert reconstruct_shamir(output_shares[:2]) != np.prod(nums)
```