# CS295/395: Secure Distributed Computation
## In-Class Exercise, 10/23/2020

## Useful Definitions for FV FHE Scheme

Reference: [Somewhat Practical Fully Homomorphic Encryption](https://eprint.iacr.org/2012/144.pdf) (2012).

```python
import random
import numpy as np

#q = 2**13 - 1
q = 2**16
p = q**3

t = 2

def noise():
    return random.randint(-5, 5) % q

def keygen():
    s = random.randint(0, t-1)
    a = random.randint(1, q-1)
    e = noise()
    pk = (-(a*s+e)%q, a)
    return s,pk

def eval_keygen(sk):
    s = sk
    a = random.randint(1, p*q-1)
    e = noise()
    rlk = (-(a*s + e) + p * s**2) % (p*q)
    return (rlk, a)

def encrypt(pk, m):
    p0, p1 = pk
    u = random.randint(0, t-1)
    e1 = noise()
    e2 = noise()
    Delta = int(q / t)
    ct1 = (p0*u + e1 + Delta*m) % q
    ct2 = (p1 * u + e2) % q
    return (ct1, ct2)

def decrypt(sk, ct):
    s = sk
    c0, c1 = ct
    print('Decrypt, before rounding:', t * ((c0 + c1 * s) % q) / q)
    m = round(t * ((c0 + c1 * s) % q) / q) % t
    return m
```

## Question 1

Implement addition for the somewhat homomorphic scheme of [FV12].

<!-- grade_id: cell-199cca6c84ee542e, points: 0, grade: false, solution: true -->
```python
def e_add(ct1, ct2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-02bce19dd026cd5c, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
ct1 = encrypt(pk, 1)
ct2 = encrypt(pk, 1)
ct3 = e_add(ct1, ct2)
m = decrypt(sk, ct3)
print('Addition result:', m)
assert m == 0
```

## Question 2

Implement multiplication for the somewhat homomorphic cryptosystem of FV12.

<!-- grade_id: cell-f7a1870c97744c01, points: 0, grade: false, solution: true -->
```python
def e_mul(ct1, ct2, rlk):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f2ff026faaba6ca1, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
rlk = eval_keygen(sk)
ct1 = encrypt(pk, 1)
ct2 = encrypt(pk, 1)
ct3 = e_mul(ct1, ct2, rlk)
m = decrypt(sk, ct3)
print('Multiplication result:', m)
assert m == 1
```

<!-- grade_id: cell-a066975f838dbc1c, points: 0, grade: true, solution: false -->
```python
# TEST CASE: How much noise is there?
total = encrypt(pk, 1)

for i in range(3000):
    new_bit = encrypt(pk, 1)
    total = e_mul(total, new_bit, rlk)

m = decrypt(sk, total)
print('Result:', m)
assert m == 1
```