# CS295/395: Secure Distributed Computation
## In-Class Exercise, 10/21/2020

## Useful Definitions for FV FHE Scheme

Reference: [Somewhat Practical Fully Homomorphic Encryption](https://eprint.iacr.org/2012/144.pdf) (2012).

```python
import random
import numpy as np

q = 2**16
p = q**3

t = 2

def noise():
    return random.randint(-5, 5) % q
```

## Question 1

Implement key generation for the somewhat homomorphic scheme in [FV12].

<!-- grade_id: cell-db5837ddd9e3f3dd, points: 0, grade: false, solution: true -->
```python
def keygen():
    # YOUR CODE HERE
    raise NotImplementedError()

keygen()
```

## Question 2

Implement encryption for the somewhat homomorphic scheme from [FV12].

<!-- grade_id: cell-f9ac89e2e958a188, points: 0, grade: false, solution: true -->
```python
def encrypt(pk, m):
    # YOUR CODE HERE
    raise NotImplementedError()
```

## Question 3

Implement decryption for the somewhat homomorphic scheme from [FV12].

<!-- grade_id: cell-f8a4f184e804c5fb, points: 0, grade: false, solution: true -->
```python
def decrypt(sk, ct):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b03b8f6f96180d97, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
m_orig = 1
print('Message:', m_orig)
ct = encrypt(pk, m_orig)
print('Ciphertext:', ct)
m = decrypt(sk, ct)
print('Decrypted message:', m)
assert m == m_orig
```

## Question 4

Implement addition for the somewhat homomorphic scheme of [FV12].

<!-- grade_id: cell-199cca6c84ee542e, points: 0, grade: false, solution: true -->
```python
def e_add(ct1, ct2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-02bce19dd026cd5c, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
ct1 = encrypt(pk, 1)
ct2 = encrypt(pk, 1)
ct3 = e_add(ct1, ct2)
m = decrypt(sk, ct3)
print('Addition result:', m)
assert m == 0
```