# CS295/395: Secure Distributed Computation
## In-Class Exercise, week of 10/17/2022

## References:
- [The Paillier Cryptosystem](https://en.wikipedia.org/wiki/Paillier_cryptosystem)
- [The ElGamal Cryptosystem](https://en.wikipedia.org/wiki/ElGamal_encryption)

```python
import galois
import random
bits = 128
```

## Question -2

Describe the main idea behind Yao's garbled circuit protocol.

<!-- grade_id: cell-5823e1abb2be85a3, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question -1

What are the primary advantages and disadvantages of Yao's garbled circuit protocol?

<!-- grade_id: cell-ae2c4c602f1e12d6, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 1

Implement key generation for the Paillier cryptosystem.

<!-- grade_id: cell-9cfbfca746bcf461, points: 0, grade: false, solution: true -->
```python
def keygen(bits):
    """Generates keys with `bits`-bits of security. Returns a pair: (secret key, public key)."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

## Question 2

Implement encryption for the Pallier cryptosystem.

<!-- grade_id: cell-204dfea9c3351633, points: 0, grade: false, solution: true -->
```python
def encrypt(m, pk):
    """Encrypts the message `m` with public key `pk`."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

## Question 3

Implement decryption for the Pallier cryptosystem.

<!-- grade_id: cell-fbdbb69c15a99cc5, points: 0, grade: false, solution: true -->
```python
def decrypt(c, sk, pk):
    """Decrypts the ciphertext `c` using secret key `sk` and public key `pk`."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b5a185ca2a26020f, points: 10, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen(128)
ciphertext = encrypt(25, pk)
plaintext = decrypt(ciphertext, sk, pk)
plaintext
```

## Question 4

Define addition for ciphertexts.

<!-- grade_id: cell-88de461defcc5020, points: 0, grade: false, solution: true -->
```python
def e_add(c1, c2, pk):
    """Add one encrypted integer to another"""
    n, g = pk

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-947746edfac16b56, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen(128)
ct1 = encrypt(50, pk)
ct2 = encrypt(10, pk)
ct3 = e_add(ct1, ct2, pk)

assert decrypt(ct3, sk, pk) == 60
```

## Question 5

Define addition-by-constant for ciphertexts.

<!-- grade_id: cell-3331edb24d2e0750, points: 0, grade: false, solution: true -->
```python
def e_add_const(c, m, pk):
    """Add constant n to an encrypted integer"""
    n, g = pk

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c8b3fb96cd63d1b2, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen(128)
ct1 = encrypt(50, pk)
ct3 = e_add_const(ct1, 10, pk)

assert decrypt(ct3, sk, pk) == 60
```

## Question 6

Define multiplication-by-constant for ciphertexts.

<!-- grade_id: cell-4858d470b77fcfb4, points: 0, grade: false, solution: true -->
```python
def e_mul_const(c, m, pk):
    """Multiplies an ancrypted integer by a constant"""
    n, g = pk
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# TEST CASE
sk, pk = keygen(128)
ct1 = encrypt(50, pk)
ct3 = e_mul_const(ct1, 3, pk)

assert decrypt(ct3, sk, pk) == 150
```

## Question 7

Implement key generation, encryption, and decryption for the ElGamal cryptosystem.

<!-- grade_id: cell-18ca2f5e2b7f9bec, points: 0, grade: false, solution: true -->
```python
def keygen(q):
    # Public key: (G, q, g, h)
    # (G is implicit here)
    # Private key: x

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-79992aaf2027f199, points: 0, grade: false, solution: true -->
```python
def encrypt(msg, pk):
    q, g, h = pk
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-95ebe53b4bd6cbb7, points: 0, grade: false, solution: true -->
```python
def decrypt(c, x, pk):
    c1, c2 = c
    q, g, h = pk

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b28eaa34531ddcba, points: 5, grade: true, solution: false -->
```python
# TEST CASE
q = 2**127 - 1

sk, pk = keygen(q)
c = encrypt(50, pk) 

decrypt(c, sk, pk)
```

## Question 8

Implement multiplication of ElGamal ciphertexts.

<!-- grade_id: cell-0e9ac5e36bbf07c6, points: 0, grade: false, solution: true -->
```python
def e_mult(c1, c2, pk):
    q, g, h = pk

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ca093a01b7afda2a, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen(q)
c1 = encrypt(5, pk)
c2 = encrypt(10, pk)
c3 = e_mult(c1, c2, pk)

assert decrypt(c3, sk, pk) == 50
```

## Question 9

Describe the differences between the Pallier and ElGamal cryptosystems.

<!-- grade_id: cell-a4d17fe6ac46f2a0, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE