# CS295/395: Secure Distributed Computation
## In-Class Exercise, 10/16/2020

```python
# Imports and definitions
import numpy as np
from collections import defaultdict
from collections import namedtuple
import urllib.request
```

## ElGamal Cryptosystem

Adapted from https://www.geeksforgeeks.org/elgamal-encryption-algorithm/

```python
# Python program to illustrate ElGamal encryption 
  
import random  
from math import pow
  
a = random.randint(2, 10) 

def invmod(a, p, maxiter=1000000):
    """The multiplicitive inverse of a in the integers modulo p:
         a * b == 1 mod p
       Returns b.
       (http://code.activestate.com/recipes/576737-inverse-modulo-p/)"""
    if a == 0:
        raise ValueError('0 has no inverse mod %d' % p)
    r = a
    d = 1
    for i in range(min(p, maxiter)):
        d = ((p // r + 1) * d) % p
        r = (d * a) % p
        if r == 1:
            break
    else:
        raise ValueError('%d has no inverse mod %d' % (a, p))
    return d

def gcd(a, b): 
    if a < b: 
        return gcd(b, a) 
    elif a % b == 0: 
        return b; 
    else: 
        return gcd(b, a % b) 

# Modular exponentiation 
def power(a, b, c): 
    x = 1
    y = a 
  
    while b > 0: 
        if b % 2 == 0: 
            x = (x * y) % c; 
        y = (y * y) % c 
        b = int(b / 2) 
  
    return x % c 

# Key generation, encryption, and decryption

def gen_key(q): 
    # Public key: (G, q, g, h)
    # (G is implicit here)
    # Private key: x
    g = random.randint(2, q) 
    x = random.randint(1, q)
    h = power(g, x, q) 

    pub = (q, g, h)
    priv = x
    return pub, priv

def encrypt(msg, pub):
    q, g, h = pub
    y = random.randint(1, q)

    s = power(h, y, q)
    c1 = power(g, y, q)
    c2 = msg * s
    return c1, c2
  
def decrypt(c1, c2, x, q): 
    h = power(c1, x, q)
    s_inv = invmod(h, q)
    dr_msg = c2 * s_inv % q
          
    return dr_msg 
```

```python
# Encryption and decryption

# q should be prime
q = 2**31 - 1

pub, priv = gen_key(q)
c1, c2 = encrypt(50, pub) 

decrypt(c1, c2, priv, q)
```

```python
# Encrypting twice with the same public key yields different ciphertexts
pub, priv = gen_key(q)
c1_1, c2_1 = encrypt(50, pub)
c1_2, c2_2 = encrypt(50, pub)
print("First:", c1_1, c2_1)
print("Second:", c1_2, c2_2)
```

## Question 1

Implement `e_mult`, multiplication for ciphertexts.

<!-- grade_id: cell-8afc4360a11fb88c, points: 0, grade: false, solution: true -->
```python
def e_mult(ct1, ct2, q):
    c1_1, c2_1 = ct1
    c1_2, c2_2 = ct2
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-da2b1ee8bb93e240, points: 5, grade: true, solution: false -->
```python
# TEST CASE
pub, priv = gen_key(q)
ct1 = encrypt(5, pub)
ct2 = encrypt(10, pub)
ct3_1, ct3_2 = e_mult(ct1, ct2, q)

assert decrypt(ct3_1, ct3_2, priv, q) == 50
```

## Question 2

Implement `e_product`, the product of a list of ciphertexts.

<!-- grade_id: cell-b16af3c8d8d9e5cd, points: 0, grade: false, solution: true -->
```python
def e_product(cts, q):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a5db87cc50a8bb5f, points: 5, grade: true, solution: false -->
```python
# TEST CASE
pub, priv = gen_key(q)

numbers = [random.randint(1, 5) for _ in range(10)]
print('List of numbers:', numbers)
cts = [encrypt(n, pub) for n in numbers]

product_1, product_2 = e_product(cts, q)
r = decrypt(product_1, product_2, priv, q)
print('Decrypted result:', r)
print('Actual product:', np.prod(numbers))

assert r == np.prod(numbers)
```