# CS295/395: Secure Distributed Computation
## In-Class Exercise, Week of 9/12/2022

```python
# Imports and definitions
import numpy as np
from collections import defaultdict
import galois

GF = galois.GF(97)
```

## Question 1

Write code to generate shares of a secret $x$ in a $(t, n)$-secret sharing scheme using Shamir's technique, where $n = 5$ and $t = 2$.

<!-- grade_id: cell-5009ee71e62855fc, points: 0, grade: false, solution: true -->
```python
def shr_2_5(v):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Example for question 1

shr_2_5(GF(5))
```

<!-- grade_id: cell-555fb4fef5fb1a22, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1

assert len(shr_2_5(GF(5))) == 5
```

## Question 2

Write a function to reconstruct the secret, using only two shares.

<!-- grade_id: cell-e32dd909f78b8f42, points: 0, grade: false, solution: true -->
```python
def reconstruct(s1, s2):
    # YOUR CODE HERE
    raise NotImplementedError()

shares = shr_2_5(GF(5))
reconstruct(shares[0], shares[1])
```

<!-- grade_id: cell-65162c80c83feaa7, points: 5, grade: true, solution: false -->
```python
# TEST CASE
shares = shr_2_5(GF(5))
assert reconstruct(shares[0], shares[1]) == GF(5)
```

## Question 3

Why is a threshold secret sharing scheme more useful than the simpler additive secret sharing scheme we saw earlier?

<!-- grade_id: cell-021796358da78ff7, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4

Write code to generate shares of a secret $x$ in a $(t, n)$-secret sharing scheme using Shamir's technique, for any $t$ and $n$.

<!-- grade_id: cell-34c3029e58f042c0, points: 0, grade: false, solution: true -->
```python
def shamir_share(v, t, n):
    # YOUR CODE HERE
    raise NotImplementedError()

shamir_share(GF(5), 3, 6)
```

```python
# Example for question 1

shr_2_5(GF(5))
```

<!-- grade_id: cell-a6374001a42a98cf, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert len(shamir_share(GF(5), 3, 6)) == 6
shares = shamir_share(GF(5), 2, 6)
assert reconstruct(shares[0], shares[1]) == GF(5)
```

## Question 5

Given the two sets of shares `shares1` and `shares2` below, write a function whose output is their sum (as a set of shares).

<!-- grade_id: cell-96a3bb1415b78611, points: 0, grade: false, solution: true -->
```python
shares1 = shamir_share(GF(20), 2, 6)
shares2 = shamir_share(GF(5), 2, 6)

def add_shares(shares1, shares2):
    # YOUR CODE HERE
    raise NotImplementedError()

added_shares = add_shares(shares1, shares2)
print(added_shares)
reconstruct(added_shares[0], added_shares[1])
```

<!-- grade_id: cell-fc07ee60a24605f8, points: 5, grade: true, solution: false -->
```python
# TEST CASE
added_shares = add_shares(shares1, shares2)
assert reconstruct(added_shares[0], added_shares[2]) == GF(25)
```

## Question 6

Write a function to reconstruct a secret from a set of at least $t$ shares. Use the `galois.lagrange_poly` function, which implements [Lagrange interpolation](https://en.wikipedia.org/wiki/Lagrange_polynomial).

<!-- grade_id: cell-8e6f22549d0766db, points: 0, grade: false, solution: true -->
```python
def reconstruct(shares):
    # YOUR CODE HERE
    raise NotImplementedError()

reconstruct(added_shares)
```

<!-- grade_id: cell-0fc28d5b0ecdae18, points: 5, grade: true, solution: false -->
```python
# TEST CASE
shares = shamir_share(GF(30), 5, 10)
assert reconstruct(shares) == GF(30)
assert reconstruct(shares[:5]) == GF(30)  # t shares are sufficient
assert reconstruct(shares[:4]) != GF(30)  # t - 1 shares are not sufficient
```

## Question 7

Given the two sets of shares `shares1` and `shares2` below, write a function whose output is their product (as a set of shares).

<!-- grade_id: cell-27b6a75239e0abfc, points: 0, grade: false, solution: true -->
```python
shares1 = shamir_share(GF(20), 3, 6)
shares2 = shamir_share(GF(3), 3, 6)

def mult_shares(shares1, shares2):
    # YOUR CODE HERE
    raise NotImplementedError()

product_shares = mult_shares(shares1, shares2)
print(product_shares)
reconstruct(shares1)
reconstruct(product_shares)
```

<!-- grade_id: cell-a208fe44dcacec08, points: 5, grade: true, solution: false -->
```python
# TEST CASE
product_shares = mult_shares(shares1, shares2)

assert reconstruct(product_shares) == GF(60)
assert reconstruct(product_shares[:4]) != GF(60)  # t shares are no longer sufficient
```