# CS295/395: Secure Distributed Computation
## In-Class Exercise, Week of 8/29/2022

# Question 1

Write functions `enc` and `dec` for an additively homomorphic cryptosystem over the integers.

<!-- grade_id: cell-16a7ad376675c884, points: 0, grade: false, solution: true -->
```python
import numpy as np

def enc(x):
    # YOUR CODE HERE
    raise NotImplementedError()

def dec(pair):
    c, k = pair
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-03410747a69767fb, points: 5, grade: true, solution: false -->
```python
assert dec(enc(5)) == 5
assert dec(enc(0)) == 0
assert dec(enc(1)) == 1

c, k = enc(5)
assert dec((c + 10, k)) == 15

c2, k2 = enc(20)
assert dec((c + c2, k + k2)) == 25
```

# Question 2

Why is the above solution considered a *one-time pad*?

<!-- grade_id: cell-2307aecc419781d8, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

# Question 3

What is problematic (for security) about the use of integers in the above solution?

<!-- grade_id: cell-52c5ebcaef37ac22, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4

Write functions `plusFE` and `multFE` that add and multiply two field elements in $GF(p)$, respectively.

<!-- grade_id: cell-c9adfe88526cc324, points: 0, grade: false, solution: true -->
```python
def plusFE(p, a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def multFE(p, a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d1dfca2e7cebbcfb, points: 5, grade: true, solution: false -->
```python
assert plusFE(100, 2, 3) == 5
assert plusFE(3, 2, 3) == 2
assert plusFE(2, 1, 1) == 0
assert multFE(100, 2, 5) == 10
assert multFE(3, 2, 5) == 1
assert multFE(2, 1, 1) == 1
assert multFE(2, 1, 0) == 0
```

## Question 5

Write functions `enc` and `dec` for an additively homomorphic cryptosystem in the finite field $GF(p)$.

<!-- grade_id: cell-7e41c12148039b16, points: 0, grade: false, solution: true -->
```python
def enc(p, x):
    # YOUR CODE HERE
    raise NotImplementedError()

def dec(p, pair):
    c, k = pair
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-94eb8fd487f35605, points: 5, grade: true, solution: false -->
```python
assert dec(100, enc(100, 5)) == 5
assert dec(2, enc(2, 0)) == 0
assert dec(2, enc(2, 1)) == 1

c, k = enc(100, 5)
assert dec(100, (c + 10, k)) == 15
```

## Question 6

Write a function `share` that generates $n$ additive secret shares from a single field element $x \in GF(p)$.

<!-- grade_id: cell-fcf3600eacf9c717, points: 0, grade: false, solution: true -->
```python
def sumFE(p, xs):
    total = 0
    for x in xs:
        total = plusFE(p, x, total)
    return total

def share(p, n, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-79cd6b0e56a18701, points: 5, grade: true, solution: false -->
```python
assert sumFE(100, share(100, 2, 5)) == 5
assert sumFE(2, share(2, 5, 1)) == 1
assert sumFE(2, share(2, 5, 0)) == 0
assert sumFE(2, share(2, 50, 1)) == 1
```

```python
# Example
print(share(100, 2, 5))
print(share(2, 5, 1))
```