# CS295/395: Secure Distributed Computation
## In-Class Exercise, week of 10/24/2022

## Useful Definitions for FV FHE Scheme

Reference: [Somewhat Practical Fully Homomorphic Encryption](https://eprint.iacr.org/2012/144.pdf) (2012).

```python
import random
import numpy as np

q = 2**16
p = q**3

t = 2

def noise():
    return random.randint(-5, 5) % q
```

## Question 1

Implement key generation for the somewhat homomorphic scheme in [FV12].

<!-- grade_id: cell-db5837ddd9e3f3dd, points: 0, grade: false, solution: true -->
```python
def keygen():
    # YOUR CODE HERE
    raise NotImplementedError()

keygen()
```

## Question 2

Implement encryption for the somewhat homomorphic scheme from [FV12].

<!-- grade_id: cell-f9ac89e2e958a188, points: 0, grade: false, solution: true -->
```python
def encrypt(pk, m):
    # YOUR CODE HERE
    raise NotImplementedError()
```

## Question 3

Implement decryption for the somewhat homomorphic scheme from [FV12].

<!-- grade_id: cell-f8a4f184e804c5fb, points: 0, grade: false, solution: true -->
```python
def decrypt(sk, ct):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b03b8f6f96180d97, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
m_orig = 1
print('Message:', m_orig)
ct = encrypt(pk, m_orig)
print('Ciphertext:', ct)
m = decrypt(sk, ct)
print('Decrypted message:', m)
assert m == m_orig
```

## Question 4

Implement addition for the somewhat homomorphic scheme of [FV12].

<!-- grade_id: cell-199cca6c84ee542e, points: 0, grade: false, solution: true -->
```python
def e_add(ct1, ct2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-02bce19dd026cd5c, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
ct1 = encrypt(pk, 1)
ct2 = encrypt(pk, 1)
ct3 = e_add(ct1, ct2)
m = decrypt(sk, ct3)
print('Addition result:', m)
assert m == 0
```

## Question 5

Implement multiplication for the somewhat homomorphic cryptosystem of FV12.

<!-- grade_id: cell-f7a1870c97744c01, points: 0, grade: false, solution: true -->
```python
def eval_keygen(sk):
    # YOUR CODE HERE
    raise NotImplementedError()

def e_mul(ct1, ct2, rlk):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f2ff026faaba6ca1, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
rlk = eval_keygen(sk)
ct1 = encrypt(pk, 1)
ct2 = encrypt(pk, 1)
ct3 = e_mul(ct1, ct2, rlk)
m = decrypt(sk, ct3)
print('Multiplication result:', m)
assert m == 1
```

<!-- grade_id: cell-a066975f838dbc1c, points: 0, grade: true, solution: false -->
```python
# TEST CASE: How much noise is there?
MULTS = 10
TRIALS = 100

failures = 0
for _ in range(TRIALS):
    total = encrypt(pk, 1)

    for i in range(MULTS):
        new_bit = encrypt(pk, 1)
        total = e_mul(total, new_bit, rlk)

    m = decrypt(sk, total)
    if m != 1:
        failures += 1
print('Percentage failed results:', failures / TRIALS * 100.0)
```