# CS295/395: Secure Distributed Computation
## In-Class Exercise, 9/30/2020

```python
# Imports and definitions
import numpy as np
from collections import defaultdict
from collections import namedtuple
import urllib.request
```

```python
AddGate = namedtuple('AddGate', ['in1', 'in2'])
MultGate = namedtuple('MultGate', ['in1', 'in2'])
```

## Question 1

Write a function `sum_circuit` that builds a circuit for summing up a list of input wires.

<!-- grade_id: cell-df91e375c98fd148, points: 0, grade: false, solution: true -->
```python
def sum_circuit(wire_names):
    first_wire, *rest_wires = wire_names
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
sum_circuit([f'w{i}' for i in range(3)])
```

<!-- grade_id: cell-4694c53c288d45f2, points: 5, grade: true, solution: false -->
```python
# TEST CASE

assert sum_circuit([f'w{i}' for i in range(3)]) == \
  AddGate(in1=AddGate(in1=AddGate(in1='w0', in2='w0'), in2='w1'), in2='w2')
```

## Question 2

Write a function `half_adder` to generate a [half-adder circuit](https://en.wikipedia.org/wiki/Adder_(electronics)).

<!-- grade_id: cell-759551c02273e1a2, points: 0, grade: false, solution: true -->
```python
def half_adder(b1, b2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
half_adder('w1', 'w2')
```

<!-- grade_id: cell-7c7c90cb532c53a4, points: 5, grade: true, solution: false -->
```python
# TEST CASE

assert half_adder('w1', 'w2') == (AddGate(in1='w1', in2='w2'), MultGate(in1='w1', in2='w2'))
```

## Question 3

Write a function `full_adder` to generate a full-adder circuit.

<!-- grade_id: cell-fda873bba42b9337, points: 0, grade: false, solution: true -->
```python
def full_adder(b1, b2, cin):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
full_adder('w1', 'w2', 'c1')
```

<!-- grade_id: cell-ab738af64e085126, points: 5, grade: true, solution: false -->
```python
# TEST CASE

assert full_adder('w1', 'w2', 'c1') == \
  (AddGate(in1=AddGate(in1='w1', in2='w2'), in2='c1'),
   MultGate(in1=MultGate(in1='w1', in2='w2'), 
            in2=MultGate(in1=AddGate(in1='w1', in2='w2'), in2='c1')))
```

## Question 4

Write a function `bitstring_adder` that returns a circuit to add two binary numbers.

<!-- grade_id: cell-3a3bf300d92b4397, points: 0, grade: false, solution: true -->
```python
def bitstring_adder(bits1, bits2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
bitstring_adder(['x1', 'x2'], ['y1', 'y2'])
```

<!-- grade_id: cell-ee9210811ed5be06, points: 5, grade: true, solution: false -->
```python
# TEST CASE

result = bitstring_adder(['x1', 'x2'], ['y1', 'y2'])
expected_result = [AddGate(in1=AddGate(in1='x1', in2='y1'), in2='zero'),
                   AddGate(in1=AddGate(in1='x2', in2='y2'), 
                           in2=MultGate(in1=MultGate(in1='x1', in2='y1'), 
                                        in2=MultGate(in1=AddGate(in1='x1', in2='y1'), 
                                                     in2='zero')))]

assert result == expected_result
```