# CS295/395: Secure Distributed Computation
## In-Class Exercise, 9/2/2020

## Question 1 (5 points)

Write functions `plusFE` and `multFE` that add and multiply two field elements in $GF(p)$, respectively.

<!-- grade_id: cell-c9adfe88526cc324, points: 0, grade: false, solution: true -->
```python
def plusFE(p, a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def multFE(p, a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d1dfca2e7cebbcfb, points: 5, grade: true, solution: false -->
```python
assert plusFE(100, 2, 3) == 5
assert plusFE(3, 2, 3) == 2
assert plusFE(2, 1, 1) == 0
assert multFE(100, 2, 5) == 10
assert multFE(3, 2, 5) == 1
assert multFE(2, 1, 1) == 1
assert multFE(2, 1, 0) == 0
```

## Question 2 (5 points)

Write functions `enc` and `dec` for an additively homomorphic cryptosystem in the finite field $GF(p)$.

<!-- grade_id: cell-7e41c12148039b16, points: 0, grade: false, solution: true -->
```python
import numpy as np

def enc(p, x):
    # YOUR CODE HERE
    raise NotImplementedError()

def dec(p, pair):
    c, k = pair
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-94eb8fd487f35605, points: 5, grade: true, solution: false -->
```python
assert dec(100, enc(100, 5)) == 5
assert dec(2, enc(2, 0)) == 0
assert dec(2, enc(2, 1)) == 1

c, k = enc(100, 5)
assert dec(100, (c + 10, k)) == 15
```

## Question 3 (5 points)

Write a function `share` that generates $n$ additive secret shares from a single field element $x \in GF(p)$.

<!-- grade_id: cell-fcf3600eacf9c717, points: 0, grade: false, solution: true -->
```python
def sumFE(p, xs):
    total = 0
    for x in xs:
        total = plusFE(p, x, total)
    return total

def share(p, n, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-79cd6b0e56a18701, points: 5, grade: true, solution: false -->
```python
assert sumFE(100, share(100, 2, 5)) == 5
assert sumFE(2, share(2, 5, 1)) == 1
assert sumFE(2, share(2, 5, 0)) == 0
assert sumFE(2, share(2, 50, 1)) == 1
```

```python
# Example
print(share(100, 2, 5))
print(share(2, 5, 1))
```

```python

```