# CS295/395: Secure Distributed Computation
## In-Class Exercise, 9/16/2020

```python
# Imports and definitions
import numpy as np
from collections import defaultdict

def plusFE(p, a, b):
    """Add field elements a and b in GF(p)"""
    return (a + b) % p
    
def multFE(p, a, b):
    """Multiply field elements a and b in GF(p)"""
    return (a * b) % p
```

## Question 1

Write code to generate shares of a secret $x$ in a $(t, n)$-secret sharing scheme using Shamir's technique, where $n = 5$ and $t = 2$.

<!-- grade_id: cell-5009ee71e62855fc, points: 0, grade: false, solution: true -->
```python
def shr_2_5(p, v):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Example for question 1

shr_2_5(100, 5)

# EXAMPLE OUTPUT: [(11, 90), (3, 10), (58, 35), (34, 95), (24, 45)]
```

<!-- grade_id: cell-555fb4fef5fb1a22, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1

assert len(shr_2_5(100, 5)) == 5
```