# CS295/395: Secure Distributed Computation
## In-Class Exercise, 9/14/2020

```python
# Imports and definitions
import numpy as np
from collections import defaultdict

def plusFE(p, a, b):
    """Add field elements a and b in GF(p)"""
    return (a + b) % p
    
def multFE(p, a, b):
    """Multiply field elements a and b in GF(p)"""
    return (a * b) % p

def sumFE(p, xs):
    """Sum up a list of field elements xs in GF(p)"""
    total = 0
    for x in xs:
        total = plusFE(p, x, total)
    return total

class Party:
    """A participant in a multiparty computation protocol."""
    def __init__(self, field_size):
        """Initialize the field size and dictionary to hold received messages."""
        self.field_size = field_size
        self.received = defaultdict(list)
    
    def send(self, other, round, msg):
        """Simulate sending a message `msg` to another party `other` during round `round`"""
        other.received[round].append(msg)
```

## Question 1

Implement the `InsecureAggregationParty` class, which adds up the parties' inputs *insecurely*.

<!-- grade_id: cell-253770a6a53e06e8, points: 0, grade: false, solution: true -->
```python
class InsecureAggregationParty(Party):
    def round1(self, parties, input):
        self.input = input
        
        # YOUR CODE HERE
        raise NotImplementedError()

    def round2(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def get_view(self):
        return (self.input, self.output, self.received[1])
```

<!-- grade_id: cell-41c0a04587798b8c, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1

# field size 100, 5 parties
p = 100
parties = [InsecureAggregationParty(p) for _ in range(5)]

# run round 1
for party in parties:
    party.round1(parties, 5)

# run round 2 and output
for party in parties:
    party.round2()
    #print('Party output:', party.output)
    assert party.output == 25
```

## Question 2

Implement a simulator for the `InsecureAggregationParty` protocol, which constructs views indistinguishable from the real-world views of the protocol using only its inputs and outputs.

<!-- grade_id: cell-a9f12612df1fdec4, points: 0, grade: false, solution: true -->
```python
def simulator(n, input, output):
    """Simulates a real-world view in the ideal world. Outputs a 
    3-tuple: (input, output, received messages from round 1)"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1feb084cc713da72, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 2

# the real-world view of each party should be indistinguishable
# from the simulated view in the ideal world
# there is no randomness here, so they are really equal!
for party in parties:
    assert party.get_view() == simulator(5, 5, 25)
```