# CS295/395: Secure Distributed Computation
## In-Class Exercise, 9/23/2020

```python
# Imports and definitions
import numpy as np
from collections import defaultdict
import urllib.request

_PRIME = 2 ** 13 - 1

shamir_lib_url = "https://raw.githubusercontent.com/jnear/cs295-secure-computation/master/utils/shamir.py"

### DANGER: this line is dangerous. Make sure the URL above is correct, and has correct code.
exec(urllib.request.urlopen(shamir_lib_url).read())

def share_shamir(t, n, x, prime=_PRIME):
    shares_with_x = share_input(x, minimum=t, shares=n, prime=prime)
    return [y for x,y in shares_with_x]

def reconstruct_shamir(shares, prime=_PRIME):
    shares_with_x = list(zip(range(1, len(shares)+1), shares))
    return recover_secret(shares_with_x, prime=prime)

class Party:
    """A participant in a multiparty computation protocol."""
    def __init__(self, field_size=_PRIME):
        """Initialize the field size and dictionary to hold received messages."""
        self.field_size = field_size
        self.input = None
        self.output = None
        self.received = defaultdict(list)
    
    def send(self, other, round, msg):
        """Simulate sending a message `msg` to another party `other` during round `round`"""
        other.received[round].append(msg)

    def get_view(self):
        """Returns the view of this party: its input, output, and received messages."""
        return (self.input, self.output, dict(self.received))
```

## Question 1

Given the two sets of shares `shares1` and `shares2` below, write a function whose output is their **product** (as a set of shares).

<!-- grade_id: cell-c5e379f74efe3722, points: 0, grade: false, solution: true -->
```python
t = 3
n = 6

shares1 = share_shamir(t, n, 20)
shares2 = share_shamir(t, n, 10)

def mult_shares(shares1, shares2):
    # YOUR CODE HERE
    raise NotImplementedError()

prod_shares = mult_shares(shares1, shares2)
print('True answer:', 20 * 10)
print('Shares:', prod_shares)

print('Recovered result, 3 shares:', reconstruct_shamir(prod_shares[:3]))
print('Recovered result, 4 shares:', reconstruct_shamir(prod_shares[:4]))
print('Recovered result, 5 shares:', reconstruct_shamir(prod_shares[:5]))
print('Recovered result, 6 shares:', reconstruct_shamir(prod_shares))
```

<!-- grade_id: cell-f4e655818083ca52, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 3
assert reconstruct_shamir(mult_shares(shares1, shares2)) == 200
```

### Question 2

Write an English description of a protocol to sum $n$ numbers securely, using Shamir secret shares, in the same format as the protocol definition from our exercise on 9/11. The ideal functionality is as follows:

\begin{equation*}
\textbf{Functionality: Sum}\\
\fbox{$\mathcal{F}_{sum}(x_1, \dots, x_n) = \sum_{i=1}^n x_i$}
\end{equation*}


<!-- grade_id: cell-4196afcaacf85339, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### Question 1(b)

Implement your protocol from question 1(a), using Shamir secret sharing.

<!-- grade_id: cell-e502961900480253, points: 0, grade: false, solution: true -->
```python
class SumParty(Party):
    def round1(self, parties, input_num):
        self.input = input_num
        self.parties = parties
        n = len(parties)
        t = n-1

        # YOUR CODE HERE
        raise NotImplementedError()

    def round2(self):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def round3(self):
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-a184246677396fcd, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1(b)

NUM_PARTIES = 10

parties = [SumParty() for _ in range(NUM_PARTIES)]

nums = [np.random.randint(0, 10) for _ in range(NUM_PARTIES)]

for p,n in zip(parties, nums):
    p.round1(parties, n)
for p in parties:
    p.round2()
for p in parties:
    p.round3()
for p in parties:
    print(p.get_view())
    assert p.output == np.sum(nums)
```