# CS295/395: Secure Distributed Computation
## In-Class Exercise, 9/4/2020

## Question 1 (5 points)

Write a function `share` that generates $n$ additive secret shares from a single field element $x \in GF(p)$.

<!-- grade_id: cell-fcf3600eacf9c717, points: 0, grade: false, solution: true -->
```python
import numpy as np

def plusFE(p, a, b):
    return (a + b) % p
    
def multFE(p, a, b):
    return (a * b) % p

def sumFE(p, xs):
    total = 0
    for x in xs:
        total = plusFE(p, x, total)
    return total

def share(p, n, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-79cd6b0e56a18701, points: 5, grade: true, solution: false -->
```python
# Test cases for automatic grading
assert sumFE(100, share(100, 2, 5)) == 5
assert sumFE(2, share(2, 5, 1)) == 1
assert sumFE(2, share(2, 5, 0)) == 0
assert sumFE(2, share(2, 50, 1)) == 1
```

```python
# Example
print(share(100, 2, 5))
print(share(2, 5,dsfsf1))
```

## Question 2 (3 points)

Is the additive secret sharing scheme for MPC secure against a *semi-honest* adversary? If not, why not?

<!-- grade_id: cell-7867ffd80feb6347, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (3 points)

Is the additive secret sharing scheme for MPC secure against a *malicious* adversary? If not, why not?

<!-- grade_id: cell-07e156948f28d61b, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE